**Important notes:**

**Important notes:**

- You can complete this exercise in 1) Colab or 2) on your local machine
  1. **Colab**: Click on the 🚀 symbol at the top of the page and select Colab. When you finished the exercise, download the file: `File` > `Download` > `ipynb`. 
  2. **Local**: click on the download button at the top of the page and choose `.ipynb`. Activate the conda environment `mr` before you start: `conda activate mr`


- Don't change the name of the file and don't delete any cells.


- Make sure you fill in any place that says  <font color='green'> \# YOUR CODE HERE </font> or "YOUR ANSWER HERE", as well as your name and (if necessary) collaborators below.


- The function **NotImplementedError()** prevents you from hand in assignments with empty cells. Simply delete the function if you start working on a cell with this entry.


- Before you turn this problem in (i.e., after you completed all tasks), make sure everything runs as expected: Restart the kernel and run all cells:
  - in *Colab*: in the menubar, select `Runtime` and click on `Restart and run all`
  - if you use *Jupyter Notebook*: in the menubar, select `Kernel` and click on `Restart & Run All`
  - if you use *Visual Studio Code*: select "Restart" and then "Run All" 


Good luck!

```python
NAME = ""
COLLABORATORS = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Mean squared error

## Setup

We need the following modules:

- pandas
- scikit-learn

```python
import pandas as pd
from sklearn.metrics import mean_squared_error
```

## Data

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

### Data structure

```python
df
```

```python
df.info()
```

## Model

```python
intercept = 500 
slope = 2

df['sales_prediction'] = intercept + slope * df['ads'] 
```

### Calculate mean squared error

---

Step 1: calculate the squared difference between true outcome and prediction:


The general formula:


$$Squared\ errors \ (SE) = (y_i-\hat{y_i})^2$$



For our data this means:

$$ se = (sales_i - sales\_prediction_i)^2$$

Hint

---

```python
df['___'] = (df['___'] - df['___']) ___ ___

```

---

- Use the `**` (power) operator to raise a value to the power of 2.

<!-- grade_id: se, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df.head()
```

<!-- grade_id: c-se, points: 2, grade: true, solution: false -->
```python
# Check your code
assert df.loc[0, 'se'] == 40000
```

---

Step 2: Sum up all squared errors

Calculate the sum of squared errors (SSE)

The general formula:

$$Sum\ of\ squared\ errors \ (SSE) = \sum_{i=1}^{n}(y_i-\hat{y_i})^2$$

which equals

$$Sum\ of\ squared\ errors \ (SSE) = \sum_{i=1}^{n} SE$$



For our data this means:



$$sse = \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2$$


which equals

$$sse = \sum_{i=1}^{n} se$$

Hint

Sum up the squared errors and save the result in a new object called `sse` (not in the Pandas dataframe)

---

```python
___ = df['___'].___
```

---

- Save the result as `sse`
- use the method `.sum()` to sum up all `se`

<!-- grade_id: sse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
sse
```

<!-- grade_id: c-sse, points: 2, grade: true, solution: false -->
```python
# Chack your code
assert sse == 14520000
```

---

Step 3: Divide by the number of examples

The general formula:

$$Mean\ squared\ error \ (MSE) = \frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})^2$$

which equals

$$Mean\ squared\ error \ (MSE) = \frac{1}{n} \times SSE$$

*with n = the number of observations*



For our data this means:



$$mse = \frac{1}{n} \times \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2$$

which equals

$$mse = \frac{1}{n} \times sse $$


*with n = the number of observations*



```python
# Obtain number of examples (observations)
n = len(df)
```

```python
n
```

Calculate the mse and save the result as `mse`:

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
mse
```

```python
# alternative way to print the result
print(f'The MSE equals {mse:.0f}')
```

<!-- grade_id: c-mse, points: 2, grade: true, solution: false -->
```python
# Check your code
assert 1451000 < mse < 1453000
```

### Root mean squared error

---

Another important measure is the root mean squared error:

The general formula:

$$Root\ mean\ squared\ error \ (RMSE) = \sqrt{\frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})^2}$$

which equals

$$Root\ mean\ squared\ error \ (RMSE) = \sqrt{\frac{1}{n} \times SSE}$$

*with n = the number of observations*






For our data this means:


$$rmse = \sqrt{\frac{1}{n} \times \sum_{i=1}^{n}(sales_i-\hat{sales\_prediction_i})^2}$$

which equals

$$rmse = \sqrt{\frac{1}{n} \times sse} $$

which also equals

$$rmse = (\frac{1}{n} \times sse)^\frac{1}{2} $$



with n = the number of observations

Calculate the rmse and save the result as `rmse` (not in your Dataframe):

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
rmse
```

<!-- grade_id: c-rmse, points: 2, grade: true, solution: false -->
```python
# Check your code
assert 1203 < rmse < 1206
```

### Mean absolute error

Another performance measure (which is less commonly used) is the mean absolute error:

The general formula:

$$Mean\ absolute\ error \ (MAE) = \frac{1}{n} \times \sum_{i=1}^{n}(y_i-\hat{y_i})$$



### Use functions

#### Mean squared error

```python
mean_squared_error(df['sales'], df['sales_prediction'])
```

#### Root mean squared error

```python
mean_squared_error(df['sales'], df['sales_prediction'], squared=False)
```

#### Mean absolute error

```python
mean_absolute_error(df['sales'], df['sales_prediction'])
```